---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`
``python
var = a + b;

`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: MultipleChoiceExercise
key: 084fb1f09e
xp: 50
```



`@possible_answers`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: PureMultipleChoiceExercise
key: 07555e3ad7
xp: 50
```



`@hint`


`@possible_answers`


`@feedback`


---

## Insert exercise title here

```yaml
type: TabExercise
key: c89626bf3b
xp: 100
```



`@pre_exercise_code`
```{python}

```

***

```yaml
type: NormalExercise
key: 7b7bad9b70
xp: 50
```

`@instructions`


`@hint`


`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

***

```yaml
type: NormalExercise
key: dbc1b375c3
xp: 50
```

`@instructions`


`@hint`


`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: BulletExercise
key: 8c1f78615f
xp: 100
```



`@pre_exercise_code`
```{python}

```

***

```yaml
type: NormalExercise
key: 0f1c895523
xp: 50
```

`@instructions`


`@hint`


`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

***

```yaml
type: NormalExercise
key: 30199a7add
xp: 50
```

`@instructions`


`@hint`


`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: VideoExercise
key: 7d7ad80885
xp: 50
```

`@projector_key`
8dca788441b26c5012a4ed895770cec9

---

## Insert exercise title here

```yaml
type: MultipleChoiceExercise
key: 340cd1f1fe
xp: 50
```



`@possible_answers`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: PureMultipleChoiceExercise
key: 2e329e908a
xp: 50
```



`@hint`


`@possible_answers`


`@feedback`


---

## Insert exercise title here

```yaml
type: VideoExercise
key: 8953bad58c
xp: 50
```

`@projector_key`
5afa99942a10decd5bdff4a8ac8a4c2b
